# Assignment 3
All questions are weighted the same in this assignment. This assignment requires more individual learning then the last one did - you are encouraged to check out the [pandas documentation](http://pandas.pydata.org/pandas-docs/stable/) to find functions or methods you might not have used yet, or ask questions on [Stack Overflow](http://stackoverflow.com/) and tag them as pandas and python related. All questions are worth the same number of points except question 1 which is worth 17% of the assignment grade.

**Note**: Questions 2-13 rely on your question 1 answer.

```python
import pandas as pd
import numpy as np
```

### Question 1
Load the energy data from the file `assets/Energy Indicators.xls`, which is a list of indicators of [energy supply and renewable electricity production](assets/Energy%20Indicators.xls) from the [United Nations](http://unstats.un.org/unsd/environment/excel_file_tables/2013/Energy%20Indicators.xls) for the year 2013, and should be put into a DataFrame with the variable name of **Energy**.

Keep in mind that this is an Excel file, and not a comma separated values file. Also, make sure to exclude the footer and header information from the datafile. The first two columns are unneccessary, so you should get rid of them, and you should change the column labels so that the columns are:

`['Country', 'Energy Supply', 'Energy Supply per Capita', '% Renewable]`

Convert `Energy Supply` to gigajoules (**Note: there are 1,000,000 gigajoules in a petajoule**). For all countries which have missing data (e.g. data with "...") make sure this is reflected as `np.NaN` values.

Rename the following list of countries (for use in later questions):

```"Republic of Korea": "South Korea",
"United States of America": "United States",
"United Kingdom of Great Britain and Northern Ireland": "United Kingdom",
"China, Hong Kong Special Administrative Region": "Hong Kong"```

There are also several countries with parenthesis in their name. Be sure to remove these, e.g. `'Bolivia (Plurinational State of)'` should be `'Bolivia'`. Additionally, there are several countries with Numeric digits in their name. Make sure to remove these as well, e.g. `'Italy9'` should be `'Italy'`. 

Next, load the GDP data from the file `assets/world_bank.csv`, which is a csv containing countries' GDP from 1960 to 2015 from [World Bank](http://data.worldbank.org/indicator/NY.GDP.MKTP.CD). Call this DataFrame **GDP**. 

Make sure to skip the header, and rename the following list of countries:

```"Korea, Rep.": "South Korea", 
"Iran, Islamic Rep.": "Iran",
"Hong Kong SAR, China": "Hong Kong"```

Finally, load the [Sciamgo Journal and Country Rank data for Energy Engineering and Power Technology](http://www.scimagojr.com/countryrank.php?category=2102) from the file `assets/scimagojr-3.xlsx`, which ranks countries based on their journal contributions in the aforementioned area. Call this DataFrame **ScimEn**.

Join the three datasets: GDP, Energy, and ScimEn into a new dataset (using the intersection of country names). Use only the last 10 years (2006-2015) of GDP data and only the top 15 countries by Scimagojr 'Rank' (Rank 1 through 15). 

The index of this DataFrame should be the name of the country, and the columns should be ['Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations',
       'Citations per document', 'H index', 'Energy Supply',
       'Energy Supply per Capita', '% Renewable', '2006', '2007', '2008',
       '2009', '2010', '2011', '2012', '2013', '2014', '2015'].

*This function should return a DataFrame with 20 columns and 15 entries, and the rows of the DataFrame should be sorted by "Rank".*

```python
!pip install xlrd
```

Output:
```
Requirement already satisfied: xlrd in c:\users\user\appdata\local\programs\python\python39\lib\site-packages (2.0.1)

[notice] A new release of pip is available: 24.0 -> 24.1.1
[notice] To update, run: python.exe -m pip install --upgrade pip
```

* remove header, skip first and last few rows, only use columns 2,3,4,5
* rename columns
* replace values w ... to NaN
* multiple values in a column by 10^6
* rename few countries 

```python
Energy = pd.read_excel('assets/Energy Indicators.xls', header = None, skiprows = 18, skipfooter = 38, usecols = [2,3,4,5])
Energy.columns = ['Country', 'Energy Supply', 'Energy Supply per Capita', '% Renewable']
Energy['Country'] = Energy['Country'].str.replace(r" \(.*\)","")

Energy['Energy Supply'] = Energy['Energy Supply'].apply(lambda x: x*1000000)

Energy['Country'].replace({"Republic of Korea": "South Korea", "United States of America20": "United States",
                           "United Kingdom of Great Britain and Northern Ireland19": "United Kingdom", "China, Hong Kong Special Administrative Region3": "Hong Kong",
                           "Bolivia (Plurinational State of)": "Bolivia", "Australia1": "Australia", "China2": "China", "Denmark5": "Denmark", "France6": "France",
                           "Greenland7": "Greenland", "Indonesia8": "Indonesia", "China, Macao Special Administrative Region4": "Macao",
                            "Italy9": "Italy", "Japan10": "Japan", "Kuwait11": "Kuwait", "Netherlands12": "Netherlands", "Portugal13": "Portugal",
                             "Saudi Arabia14": "Saudia Arabia", "Serbia15": "Serbia", "Spain16": "Spain", "Switzerland17": "Switzerland", "Ukraine18": "Ukraine"}, inplace = True)
Energy
```

Output:
```
Country  \
0                  Afghanistan   
1                      Albania   
2                      Algeria   
3               American Samoa   
4                      Andorra   
..                         ...   
222                   Viet Nam   
223  Wallis and Futuna Islands   
224                      Yemen   
225                     Zambia   
226                   Zimbabwe   

                                         Energy Supply  \
0                                            321000000   
1                                            102000000   
2                                           1959000000   
3    .................................................   
4                                              9000000   
..                                                 ...   
222                                         2554000000   
223                                                  0   
224                                          344000000   
225                               
```

```python
GDP = pd.read_csv('assets/world_bank.csv', skiprows = 4)
GDP['Country Name'].replace({"Korea, Rep.": "South Korea",  "Iran, Islamic Rep.": "Iran", "Hong Kong SAR, China": "Hong Kong"}, inplace = True)
GDP.rename(columns = {"Country Name": "Country"}, inplace = True)
GDP
```

Output:
```
Country Country Code                            Indicator Name  \
0               Aruba          ABW  GDP at market prices (constant 2010 US$)   
1             Andorra          AND  GDP at market prices (constant 2010 US$)   
2         Afghanistan          AFG  GDP at market prices (constant 2010 US$)   
3              Angola          AGO  GDP at market prices (constant 2010 US$)   
4             Albania          ALB  GDP at market prices (constant 2010 US$)   
..                ...          ...                                       ...   
259       Yemen, Rep.          YEM  GDP at market prices (constant 2010 US$)   
260      South Africa          ZAF  GDP at market prices (constant 2010 US$)   
261  Congo, Dem. Rep.          COD  GDP at market prices (constant 2010 US$)   
262            Zambia          ZMB  GDP at market prices (constant 2010 US$)   
263          Zimbabwe          ZWE  GDP at market prices (constant 2010 US$)   

     Indicator Code          1960          1961      
```

```python
ScimEn = pd.read_excel('assets/scimagojr-3.xlsx', skipfooter = 176)
ScimEn
```

Output:
```
Rank             Country  Documents  Citable documents  Citations  \
0      1               China     127050             126767     597237   
1      2       United States      96661              94747     792274   
2      3               Japan      30504              30287     223024   
3      4      United Kingdom      20944              20357     206091   
4      5  Russian Federation      18534              18301      34266   
5      6              Canada      17899              17620     215003   
6      7             Germany      17027              16831     140566   
7      8               India      15005              14841     128763   
8      9              France      13153              12973     130632   
9     10         South Korea      11983              11923     114675   
10    11               Italy      10964              10794     111850   
11    12               Spain       9428               9330     123336   
12    13                Iran       8896               8
```

```python

merge1 = pd.merge(ScimEn, Energy, how = 'inner', left_on="Country",right_on="Country")
GDP = GDP[['Country', 'Country Code', 'Indicator Name', 'Indicator Code', '2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
merge2 = pd.merge(merge1, GDP, how = 'inner', left_on="Country",right_on="Country").set_index("Country")
merge2 = merge2[['Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations', 'Citations per document', 'H index', 'Energy Supply',
                 'Energy Supply per Capita', '% Renewable', '2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
merge2
```

Output:
```
Rank  Documents  Citable documents  Citations  \
Country                                                             
China                  1     127050             126767     597237   
United States          2      96661              94747     792274   
Japan                  3      30504              30287     223024   
United Kingdom         4      20944              20357     206091   
Russian Federation     5      18534              18301      34266   
Canada                 6      17899              17620     215003   
Germany                7      17027              16831     140566   
India                  8      15005              14841     128763   
France                 9      13153              12973     130632   
South Korea           10      11983              11923     114675   
Italy                 11      10964              10794     111850   
Spain                 12       9428               9330     123336   
Australia             14       8831               8725
```

<!-- grade_id: cell-bce4d6f2ecdd1297, points: 0, grade: false, solution: true -->
```python
def answer_one():
    # YOUR CODE HERE
    Energy = pd.read_excel('assets/Energy Indicators.xls', header = None, skiprows = 18, skipfooter = 38, usecols = [2,3,4,5])
    Energy.columns = ['Country', 'Energy Supply', 'Energy Supply per Capita', '% Renewable']
    Energy['Energy Supply'] = Energy['Energy Supply'].apply(lambda x: x*1000000)

    Energy['Country'] = Energy['Country'].str.replace(r" \(.*\)","")

    Energy['Country'].replace({"Republic of Korea": "South Korea", "United States of America20": "United States",
                            "United Kingdom of Great Britain and Northern Ireland19": "United Kingdom", "China, Hong Kong Special Administrative Region3": "Hong Kong",
                            "Bolivia (Plurinational State of)": "Bolivia", "Australia1": "Australia", "China2": "China", "Denmark5": "Denmark", "France6": "France",
                            "Greenland7": "Greenland", "Indonesia8": "Indonesia", "China, Macao Special Administrative Region4": "Macao",
                                "Italy9": "Italy", "Japan10": "Japan", "Kuwait11": "Kuwait", "Netherlands12": "Netherlands", "Portugal13": "Portugal",
                                "Saudi Arabia14": "Saudia Arabia", "Serbia15": "Serbia", "Spain16": "Spain", "Switzerland17": "Switzerland", "Ukraine18": "Ukraine"}, inplace = True)
    
    GDP = pd.read_csv('assets/world_bank.csv', skiprows = 4)
    GDP['Country Name'].replace({"Korea, Rep.": "South Korea",  "Iran, Islamic Rep.": "Iran", "Hong Kong SAR, China": "Hong Kong"}, inplace = True)
    
    ScimEn = pd.read_excel('assets/scimagojr-3.xlsx', skipfooter = 176)

    merge1 = pd.merge(ScimEn, Energy, how = 'inner', left_on="Country",right_on="Country")

    GDP.rename(columns = {"Country Name": "Country"}, inplace = True)
    GDP = GDP[['Country', 'Country Code', 'Indicator Name', 'Indicator Code', '2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
    
    merge2 = pd.merge(merge1, GDP, how = 'inner', left_on="Country",right_on="Country").set_index("Country")
    merge2 = merge2[['Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations', 'Citations per document', 'H index', 'Energy Supply',
                    'Energy Supply per Capita', '% Renewable', '2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
    
    return merge1
```

### Question 2
The previous question joined three datasets then reduced this to just the top 15 entries. When you joined the datasets, but before you reduced this to the top 15 items, how many entries did you lose?

*This function should return a single number.*

```python
%%HTML
<svg width="800" height="300">
  <circle cx="150" cy="180" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="blue" />
  <circle cx="200" cy="100" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="red" />
  <circle cx="100" cy="100" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="green" />
  <line x1="150" y1="125" x2="300" y2="150" stroke="black" stroke-width="2" fill="black" stroke-dasharray="5,3"/>
  <text x="300" y="165" font-family="Verdana" font-size="35">Everything but this!</text>
</svg>
```

Output:
```
<IPython.core.display.HTML object>
```

<!-- grade_id: cell-b0c3202c000aada4, points: 0, grade: false, solution: true -->
```python
def answer_two():
    # YOUR CODE HERE
    Energy = pd.read_excel('assets/Energy Indicators.xls', header = None, skiprows = 18, skipfooter = 38, usecols = [2,3,4,5])
    Energy.columns = ['Country', 'Energy Supply', 'Energy Supply per Capita', '% Renewable']
    Energy['Energy Supply'] = Energy['Energy Supply'].apply(lambda x: x*1000000)
    
    Energy['Country'] = Energy['Country'].str.replace(r" \(.*\)","")
    
    Energy['Country'].replace({"Republic of Korea": "South Korea", "United States of America20": "United States",
                            "United Kingdom of Great Britain and Northern Ireland19": "United Kingdom", "China, Hong Kong Special Administrative Region3": "Hong Kong",
                            "Bolivia (Plurinational State of)": "Bolivia", "Australia1": "Australia", "China2": "China", "Denmark5": "Denmark", "France6": "France",
                            "Greenland7": "Greenland", "Indonesia8": "Indonesia", "China, Macao Special Administrative Region4": "Macao",
                                "Italy9": "Italy", "Japan10": "Japan", "Kuwait11": "Kuwait", "Netherlands12": "Netherlands", "Portugal13": "Portugal",
                                "Saudi Arabia14": "Saudia Arabia", "Serbia15": "Serbia", "Spain16": "Spain", "Switzerland17": "Switzerland", "Ukraine18": "Ukraine"}, inplace = True)

    GDP = pd.read_csv('assets/world_bank.csv', skiprows = 4)
    GDP['Country Name'].replace({"Korea, Rep.": "South Korea",  "Iran, Islamic Rep.": "Iran", "Hong Kong SAR, China": "Hong Kong"}, inplace = True)
    GDP.rename(columns = {"Country Name": "Country"}, inplace = True)
    
    ScimEn = pd.read_excel('assets/scimagojr-3.xlsx')
    
    # outer
    merge_outer1 = pd.merge(ScimEn, Energy, how = 'outer', left_on="Country",right_on="Country")
    merge_outer2 = pd.merge(merge_outer1, GDP, how = 'outer', left_on="Country",right_on="Country").set_index("Country")

    # inner
    GDP_inner = GDP[['Country', 'Country Code', 'Indicator Name', 'Indicator Code', '2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
    merge_inner1 = pd.merge(ScimEn, Energy, how = 'inner', left_on="Country",right_on="Country")
    merge_inner2 = pd.merge(merge_inner1, GDP_inner, how = 'inner', left_on="Country",right_on="Country").set_index("Country")

    result = len(merge_outer2) - len(merge_inner2)
    return result
```

### Question 3
What are the top 15 countries for average GDP over the last 10 years?

*This function should return a Series named `avgGDP` with 15 countries and their average GDP sorted in descending order.*

```python
q3 = merge2[['2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
avg_gdp = q3.mean(axis = 1)
avgGDP = avg_gdp.sort_values(ascending = False)
avgGDP
```

Output:
```
Country
United States         1.536434e+13
China                 6.348609e+12
Japan                 5.542208e+12
Germany               3.493025e+12
France                2.681725e+12
United Kingdom        2.487907e+12
Brazil                2.189794e+12
Italy                 2.120175e+12
India                 1.769297e+12
Canada                1.660647e+12
Russian Federation    1.565459e+12
Spain                 1.418078e+12
Australia             1.164043e+12
South Korea           1.106715e+12
dtype: float64
```

<!-- grade_id: cell-8c3d74335c0d489a, points: 0, grade: false, solution: true -->
```python
def answer_three():
    # YOUR CODE HERE
    merge2 = answer_one()
    q3 = merge2[['2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
    avg_gdp = q3.mean(axis = 1)
    avgGDP = avg_gdp.sort_values(ascending = False)
    return avgGDP
```

### Question 4
By how much had the GDP changed over the 10 year span for the country with the 6th largest average GDP?

*This function should return a single number.*

```python
q4 = merge2[['2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
uk = q4.loc["United Kingdom"]

change = uk.loc['2015'] - uk.loc['2006']
change
```

Output:
```
246702696075.3999
```

<!-- grade_id: cell-7f77d099e3e0bbee, points: 0, grade: false, solution: true -->
```python
def answer_four():
    # YOUR CODE HERE
    merge2 = answer_one()
    q4 = merge2[['2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
    uk = q4.loc["United Kingdom"]

    change = uk.loc['2015'] - uk.loc['2006']
    return change
```

### Question 5
What is the mean energy supply per capita?

*This function should return a single number.*

```python
q5 = merge2["Energy Supply per Capita"]
mean_energy = q5.mean()
mean_energy
```

Output:
```
160.35714285714286
```

<!-- grade_id: cell-58e79d558e982eef, points: 0, grade: false, solution: true -->
```python
def answer_five():
    # YOUR CODE HERE
    merge2 = answer_one()
    q5 = merge2["Energy Supply per Capita"]
    mean_energy = q5.mean()
    return mean_energy
```

### Question 6
What country has the maximum % Renewable and what is the percentage?

*This function should return a tuple with the name of the country and the percentage.*

```python
q6 = merge2['% Renewable']
max_renewable = q6.max()
max_index = q6.idxmax()
q6_ans = (max_index, max_renewable)
q6_ans
```

Output:
```
('Brazil', 69.64803)
```

<!-- grade_id: cell-b6824b78e74619f9, points: 0, grade: false, solution: true -->
```python
def answer_six():
    # YOUR CODE HERE
    merge2 = answer_one()
    q6 = merge2['% Renewable']
    max_renewable = q6.max()
    max_index = q6.idxmax()
    q6_ans = (max_index, max_renewable)
    return q6_ans
```

### Question 7
Create a new column that is the ratio of Self-Citations to Total Citations. 
What is the maximum value for this new column, and what country has the highest ratio?

*This function should return a tuple with the name of the country and the ratio.*

```python
q7 = merge2[['Self-citations', 'Citations']]
q7['ratio'] = q7['Self-citations'] / q7['Citations']
max_ratio = q7['ratio'].max()
max_index = q7['ratio'].idxmax()
q7_ans = (max_index, max_ratio)
q7_ans
```

Output:
```
C:\Users\User\AppData\Local\Temp\ipykernel_22176\1331210098.py:2: SettingWithCopyWarning: 
A value is trying to be set on a copy of a slice from a DataFrame.
Try using .loc[row_indexer,col_indexer] = value instead

See the caveats in the documentation: https://pandas.pydata.org/pandas-docs/stable/user_guide/indexing.html#returning-a-view-versus-a-copy
  q7['ratio'] = q7['Self-citations'] / q7['Citations']
('China', 0.6893126179389422)
```

<!-- grade_id: cell-a4f39737f38aa53c, points: 0, grade: false, solution: true -->
```python
def answer_seven():
    # YOUR CODE HERE
    merge2 = answer_one()
    q7 = merge2[['Self-citations', 'Citations']]
    q7['ratio'] = q7['Self-citations'] / q7['Citations']
    max_ratio = q7['ratio'].max()
    max_index = q7['ratio'].idxmax()
    q7_ans = (max_index, max_ratio)
    return q7_ans
```

### Question 8

Create a column that estimates the population using Energy Supply and Energy Supply per capita. 
What is the third most populous country according to this estimate?

*This function should return the name of the country*

```python
q8 = merge2[['Energy Supply', 'Energy Supply per Capita']]
q8['population'] = q8['Energy Supply'] / q8['Energy Supply per Capita']
popular = q8.sort_values(by = 'population', ascending = False)
popular_country = popular['population'].index[2]
popular_country
```

Output:
```
C:\Users\User\AppData\Local\Temp\ipykernel_22176\4087658320.py:2: SettingWithCopyWarning: 
A value is trying to be set on a copy of a slice from a DataFrame.
Try using .loc[row_indexer,col_indexer] = value instead

See the caveats in the documentation: https://pandas.pydata.org/pandas-docs/stable/user_guide/indexing.html#returning-a-view-versus-a-copy
  q8['population'] = q8['Energy Supply'] / q8['Energy Supply per Capita']
'United States'
```

<!-- grade_id: cell-9ca58137846b84d6, points: 0, grade: false, solution: true -->
```python
def answer_eight():
    # YOUR CODE HERE
    merge2 = answer_one()
    q8 = merge2[['Energy Supply', 'Energy Supply per Capita']]
    q8['population'] = q8['Energy Supply'] / q8['Energy Supply per Capita']
    popular = q8.sort_values(by = 'population', ascending = False)
    popular_country = popular['population'].index[2]
    return popular_country
```

### Question 9
Create a column that estimates the number of citable documents per person. 
What is the correlation between the number of citable documents per capita and the energy supply per capita? Use the `.corr()` method, (Pearson's correlation).

*This function should return a single number.*

*(Optional: Use the built-in function `plot9()` to visualize the relationship between Energy Supply per Capita vs. Citable docs per Capita)*

```python
Top15 = merge2
Top15['PopEst'] = Top15['Energy Supply'] / Top15['Energy Supply per Capita']
Top15['Citable docs per Capita'] = Top15['Citable documents'] / Top15['PopEst']
Top15['Citable docs per Capita'].corr(Top15['Energy Supply per Capita'])
```

Output:
```
0.7901935488300876
```

<!-- grade_id: cell-033679ea456bfb9d, points: 0, grade: false, solution: true -->
```python
def answer_nine():
    # YOUR CODE HERE
    Top15 = answer_one()
    Top15['PopEst'] = Top15['Energy Supply'] / Top15['Energy Supply per Capita']
    Top15['Citable docs per Capita'] = Top15['Citable documents'] / Top15['PopEst']
    corr = Top15['Citable docs per Capita'].corr(Top15['Energy Supply per Capita'])
    return corr
```

```python
answer_nine()
```

Output:
```
0.7901935488300876
```

```python
def plot9():
    import matplotlib as plt
    %matplotlib inline
    
    Top15 = answer_one()
    Top15['PopEst'] = Top15['Energy Supply'] / Top15['Energy Supply per Capita']
    Top15['Citable docs per Capita'] = Top15['Citable documents'] / Top15['PopEst']
    Top15.plot(x='Citable docs per Capita', y='Energy Supply per Capita', kind='scatter', xlim=[0, 0.0006])
```

<!-- grade_id: cell-3cb5c699065a4a20, points: 6.66, grade: true, solution: false -->
```python
assert answer_nine() >= -1. and answer_nine() <= 1., "Q9: A valid correlation should between -1 to 1!"

```

### Question 10
Create a new column with a 1 if the country's % Renewable value is at or above the median for all countries in the top 15, and a 0 if the country's % Renewable value is below the median.

*This function should return a series named `HighRenew` whose index is the country name sorted in ascending order of rank.*

```python
median_renewable = merge2['% Renewable'].median()
merge2['Median Check'] = np.where(merge2['% Renewable'] >= median_renewable, 1, 0)
HighRenew = merge2['Median Check']
HighRenew
```

Output:
```
Country
China                 1
United States         0
Japan                 0
United Kingdom        0
Russian Federation    1
Canada                1
Germany               1
India                 0
France                0
South Korea           0
Italy                 1
Spain                 1
Iran                  0
Australia             0
Brazil                1
Name: Median Check, dtype: int32
```

<!-- grade_id: cell-0fdf60e64bf1a4f9, points: 0, grade: false, solution: true -->
```python
def answer_ten():
    # YOUR CODE HERE
    merge2 = answer_one()
    median_renewable = merge2['% Renewable'].median()
    merge2['Median Check'] = np.where(merge2['% Renewable'] >= median_renewable, 1, 0)
    HighRenew = merge2['Median Check']
    return HighRenew
```

### Question 11
Use the following dictionary to group the Countries by Continent, then create a DataFrame that displays the sample size (the number of countries in each continent bin), and the sum, mean, and std deviation for the estimated population of each country.

```python
ContinentDict  = {'China':'Asia', 
                  'United States':'North America', 
                  'Japan':'Asia', 
                  'United Kingdom':'Europe', 
                  'Russian Federation':'Europe', 
                  'Canada':'North America', 
                  'Germany':'Europe', 
                  'India':'Asia',
                  'France':'Europe', 
                  'South Korea':'Asia', 
                  'Italy':'Europe', 
                  'Spain':'Europe', 
                  'Iran':'Asia',
                  'Australia':'Australia', 
                  'Brazil':'South America'}
```

*This function should return a DataFrame with index named Continent `['Asia', 'Australia', 'Europe', 'North America', 'South America']` and columns `['size', 'sum', 'mean', 'std']`*

```python
ContinentDict  = {'China':'Asia', 
                  'United States':'North America', 
                  'Japan':'Asia', 
                  'United Kingdom':'Europe', 
                  'Russian Federation':'Europe', 
                  'Canada':'North America', 
                  'Germany':'Europe', 
                  'India':'Asia',
                  'France':'Europe', 
                  'South Korea':'Asia', 
                  'Italy':'Europe', 
                  'Spain':'Europe', 
                  'Iran':'Asia',
                  'Australia':'Australia', 
                  'Brazil':'South America'}

q11 = merge2
q11['Continent'] = pd.Series(ContinentDict)
q11['Est Population'] = q11['Energy Supply'] / q11['Energy Supply per Capita']
df_11 = q11.groupby('Continent')['Est Population'].agg([np.size,np.sum, np.mean, np.std])
df_11
```

<!-- grade_id: cell-a5e0c0df27304f98, points: 0, grade: false, solution: true -->
```python
def answer_eleven():
    # YOUR CODE HERE
    ContinentDict  = {'China':'Asia', 
                  'United States':'North America', 
                  'Japan':'Asia', 
                  'United Kingdom':'Europe', 
                  'Russian Federation':'Europe', 
                  'Canada':'North America', 
                  'Germany':'Europe', 
                  'India':'Asia',
                  'France':'Europe', 
                  'South Korea':'Asia', 
                  'Italy':'Europe', 
                  'Spain':'Europe', 
                  'Iran':'Asia',
                  'Australia':'Australia', 
                  'Brazil':'South America'}
    
    q11 = answer_one()
    q11['Continent'] = pd.Series(ContinentDict)
    q11['Est Population'] = q11['Energy Supply'] / q11['Energy Supply per Capita']
    df_11 = q11.groupby('Continent')['Est Population'].agg([np.size,np.sum, np.mean, np.std])
    return df_11
```

### Question 12
Cut % Renewable into 5 bins. Group Top15 by the Continent, as well as these new % Renewable bins. How many countries are in each of these groups?

*This function should return a Series with a MultiIndex of `Continent`, then the bins for `% Renewable`. Do not include groups with no countries.*

```python
ContinentDict  = {'China':'Asia', 
                  'United States':'North America', 
                  'Japan':'Asia', 
                  'United Kingdom':'Europe', 
                  'Russian Federation':'Europe', 
                  'Canada':'North America', 
                  'Germany':'Europe', 
                  'India':'Asia',
                  'France':'Europe', 
                  'South Korea':'Asia', 
                  'Italy':'Europe', 
                  'Spain':'Europe', 
                  'Iran':'Asia',
                  'Australia':'Australia', 
                  'Brazil':'South America'}

Top15 = answer_one()
Top15['Continent'] = pd.Series(ContinentDict)
Top15['% Renewable']=pd.cut(Top15['% Renewable'],5)   # array issue

q12 = Top15.groupby(['Continent','% Renewable'], observed = True)['Continent'].agg(np.size).dropna()
q12
```

Output:
```
Series([], Name: Continent, dtype: object)
```

<!-- grade_id: cell-2ecd9a4076abd8f0, points: 0, grade: false, solution: true -->
```python
def answer_twelve():
    # YOUR CODE HERE
    ContinentDict  = {'China':'Asia', 
                  'United States':'North America', 
                  'Japan':'Asia', 
                  'United Kingdom':'Europe', 
                  'Russian Federation':'Europe', 
                  'Canada':'North America', 
                  'Germany':'Europe', 
                  'India':'Asia',
                  'France':'Europe', 
                  'South Korea':'Asia', 
                  'Italy':'Europe', 
                  'Spain':'Europe', 
                  'Iran':'Asia',
                  'Australia':'Australia', 
                  'Brazil':'South America'}

    Top15 = answer_one()
    Top15['Continent'] = pd.Series(ContinentDict)
    Top15['% Renewable']=pd.cut(Top15['% Renewable'],5)   # array issue

    q12 = Top15.groupby(['Continent','% Renewable'], observed = True)['Continent'].agg(np.size).dropna()
    return q12

```

### Question 13
Convert the Population Estimate series to a string with thousands separator (using commas). Use all significant digits (do not round the results).

e.g. 12345678.90 -> 12,345,678.90

*This function should return a series `PopEst` whose index is the country name and whose values are the population estimate string*

```python
q13 = merge2
q13['Est Population'] = q13['Energy Supply'] / q13['Energy Supply per Capita']
q13_series = q13['Est Population'].apply('{:,}'.format)
q13_series
```

Output:
```
0     1,367,645,161.2903225
1      317,615,384.61538464
2      127,409,395.97315437
3      63,870,967.741935484
4             143,500,000.0
5       35,239,864.86486486
6       80,369,696.96969697
7     1,276,730,769.2307692
8       63,837,349.39759036
9      49,805,429.864253394
10     59,908,256.880733944
11       46,443,396.2264151
12     23,316,017.316017315
13     205,915,254.23728815
Name: Est Population, dtype: object
```

<!-- grade_id: cell-58eb0ee0921d93fb, points: 0, grade: false, solution: true -->
```python
def answer_thirteen():
    # YOUR CODE HERE
    q13 = answer_one()
    q13['Est Population'] = q13['Energy Supply'] / q13['Energy Supply per Capita']
    q13_series = q13['Est Population'].apply('{:,}'.format)
    return q13_series
```

### Optional

Use the built in function `plot_optional()` to see an example visualization.

```python
def plot_optional():
    import matplotlib as plt
    %matplotlib inline
    Top15 = answer_one()
    ax = Top15.plot(x='Rank', y='% Renewable', kind='scatter', 
                    c=['#e41a1c','#377eb8','#e41a1c','#4daf4a','#4daf4a','#377eb8','#4daf4a','#e41a1c',
                       '#4daf4a','#e41a1c','#4daf4a','#4daf4a','#e41a1c','#dede00','#ff7f00'], 
                    xticks=range(1,16), s=6*Top15['2014']/10**10, alpha=.75, figsize=[16,6]);

    for i, txt in enumerate(Top15.index):
        ax.annotate(txt, [Top15['Rank'][i], Top15['% Renewable'][i]], ha='center')

    print("This is an example of a visualization that can be created to help understand the data. \
This is a bubble chart showing % Renewable vs. Rank. The size of the bubble corresponds to the countries' \
2014 GDP, and the color corresponds to the continent.")
```